# COMP SCI 2009 PIT W01 Practical

<!-- BEGIN QUESTION -->

## Exercise 1: Binary Search (25%)

In *Introduction to Applied Programming*, you coded a number guessing game in one of your practical assignments. In the assignment, the user was the one guessing the number, while the program told you if your guess is larger or smaller than the answer. In this exercise, you will do the opposite, i.e., the user inputs a number (between 0 ~ 999) and a number guessing AI continuously guesses the number (using a binary search method) according to the user response.

**Example (<font color=red>red text</font> represents user input):**
<pre>
Please input a number (0~999): <font color=red>523</font>
Is the number 500? <font color=red>larger</font>
Is the number 725? <font color=red>smaller</font>
Is the number 612? <font color=red>smaller</font>
...
Is the number 523? <font color=red>yes</font>
</pre>

```python
# the range of the number
left = 0
right = 999

# user input the answer
answer = int(input('Please answer a number (0~999): '))

# AI guessing loop starts
cnt = 0
while True:
    # instead of random guess, AI should use binary search
    # REPLACE the line below
    guess = (left + right) // 2
    response = input(f'Is the number {guess}?')
    
    # INSERT YOUR CODE BELOW    
    # Update left or right according to the response
    # ~ 7 lines
    if response == "yes":
        print(f"The number is {guess}.")
        break
    elif response == "larger":
        left = guess + 1
    elif response == "smaller":
        right = guess - 1
    else:
        print("Invalid response. Please enter 'yes', 'larger', or 'smaller'.")
    
    # terminate loop for autograder
    cnt += 1
    if cnt >= 15:
        break
```

Output:
```
The number is 520.
```

<!-- END QUESTION -->

## Exercise 2: Insertion sort & pseudocode (25%)

Below is the pseudocode of an **insertion sort** algorithm. Note that in this pseudocode the index starts from `1` instead of `0` and you will need to modify the code accordingly. 
```
INSERTION_SORT(A)
for j = 2 to A.length
    key = A[j]
    // Insert A[j] into the sorted sequence A[1..j-1]
    i = j - 1
    while i > 0 and A[i] > key
        A[i+1] = A[i]
        i = i - 1
    A[i+1] = key
```

Please do the following:
1. Read the pseudo code and explain how **each line of the code** works. Before you start Googling about **insertion sort**, please spend 10 minutes reading the pseudocode and make sense of it! Discuss with your tutor if you are having trouble.
2. Implement the pseudocode using Python in the cell below.

# Answer the first part of the question in this cell.

INSERTION_SORT(A) is the function declaration of the insertion sort algorithm, which takes an array A as its input and sorts it in ascending order.

for j = 2 to A.length is a loop that starts from j = 2 and goes up to the length of the array A. The variable j represents the index of the element we are currently considering for insertion into the sorted sequence.

key = A[j] stores the value of the element at index j in the variable key. This is the element that we want to insert into its correct position in the sorted subarray.

i = j - 1 initialize the variable i to j - 1. The variable i represents the index of the last element in the sorted sequence. Initially, the sorted sequence contains only the element at index 1.

while i > 0 and A [i] > key is a while loop that continues as long as i is greater than 0 (ensuring we don't go out of bounds) and the element at index i is greater than the key element. The loop is used to find the correct position in the sorted sequence to insert the key.

A[i+1] = A[i] shifts the element at index i one position to the right to make space for the insertion of the key element in the while loop. This step effectively moves larger elements one position to the right to accommodate the new element in the correct position.

i = i - 1 decrement i to move to the previous index in the sorted sequence and continue comparing with the previous element to find the correct insertion point.

A[i+1] = key means we've found the correct position to insert the key element after the while loop is finished. We place the key element at index i + 1, which is the correct position in the sorted sequence.



```python
def insert_sort(A):
    for j in range(1, len(A)):
        # INSERT YOUR CODE BELOW
        # ~6 line
        key = A[j]
        i = j - 1
        while i >= 0 and A[i] > key:
            A[i + 1] = A[i]
            i = i - 1
        A[i + 1] = key
    return A

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE

A = insert_sort([81, 14, 3, 94, 35, 31, 28, 17, 13, 86])
print(A) # [3, 13, 14, 17, 28, 31, 35, 81, 86, 94]
A = insert_sort([77, 3, 71, 25, 91, 83, 89, 69, 53, 28])
print(A) # [3, 25, 28, 53, 69, 71, 77, 83, 89, 91]
```

Output:
```
[3, 13, 14, 17, 28, 31, 35, 81, 86, 94]
[3, 25, 28, 53, 69, 71, 77, 83, 89, 91]
```

## Exercise 3: Max Consecutive Identical Characters (25%)

Complete the following function such that when it is given a string, it returns the maximum number of consecutive identical characters within that string.

**Example 1**

- Input: `aabbccccddd`<br>
- Output: 4<br>
- Comment: The input has 2 consecutive `a`, 2 `b`, 4 `c`, 3 `d`, thus the maximum number of consecutive identical characters is 4.

**Example 2**

- Input: `abcffgfffffgg`<br>
- Output: 5<br>
- Comment: Although there are 7 `f`, the longest consecutive substring is `fffff` as there is a `g` between `ff` and `fffff`. 
**Example 3**

- Input: `apple`<br>
- Output: 2<br>

```python
def consec_char(input_str):
    max_cnt = 1
    cnt = 1
    cur = input_str[0]
    
    for i in range(1, len(input_str)):
        c = input_str[i]
        # INSERT YOUR CODE BELOW
        # hint: check if c is equal to cur. If not, update cur and check max_cnt
        # <10 lines
        if c == cur:
            cnt += 1
            if cnt > max_cnt:
                max_cnt = cnt
        else:
            cnt = 1
            cur = c

    return max_cnt
# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
c = consec_char('aabbccccddd')
print(c) # 4
c = consec_char('abcffgfffffgg')
print(c) # 5
```

Output:
```
4
5
```

## Excercise 4: Longest Continuous Decreasing Subsequence (25%)

Complete the following function such that when it is given an unsorted array of integers `nums`, it returns the length of the longest continuous decreasing subsequence (i.e. subarray). The subsequence must be **strictly** decreasing.

Formally, a continuous decreasing subsequence is defined by two indices l and r (l < r) such that it is [nums[l], nums[l + 1], ..., nums[r - 1], nums[r]] and for each l <= i < r, nums[i] > nums[i + 1].

**Example 1** 
- input: `[3,5,4,3,2,7,1]`
- output: 4
- The longest continuous decreasing subsequence is `5432`. Note that although `54321` is the longest decreasing subsequence, it is not a continuous one as `7` is betweeen `2` and `1`

**Example 2**
- input: `[7,8,9,7,7,6,5,3,6]`
- output: 4
- The subsequence is [7, 6, 5, 3]. The numbers must be strictly decreasing, thus the 4th and 5th elements are both `7` and do not satisfy the requirement. 

**Example 3**
- input: `[2,4,8,6,5,8]`
- output: 3

**Example 4**
- input: `[1, 2, 3]`
- output: 1
- Special case when there is no continuous decreasing subsequence. 

```python
def lcds(array):
    max_cnt = -1
    # INSERT YOUR CODE BELOW
    current_cnt = 1
    for i in range(1, len(array)):
        if array[i] < array[i - 1]:
            current_cnt += 1
        else:
            current_cnt = 1

        if current_cnt > max_cnt:
            max_cnt = current_cnt

    return max_cnt

# =======================
# EXAMPLE TESTING CASES
# DO NOT CHANGE
c = lcds([3,5,4,3,2,7,1])
print(c) # 4
c = lcds([7,8,9,7,7,6,5,3,6])
print(c) # 4
```

Output:
```
4
4
```

Tests for Q2 (points: 0.5)
```
>>> insert_sort([-5, 10, 3])
[-5, 3, 10]
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> insert_sort([-5, 13, -5, 10])
[-5, -5, 10, 13]
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Q3 (points: 1)
```
>>> consec_char('aaaddc')
3
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> consec_char('abcwwww')
4
```
failure_message: Incorrect answer, try again!
success_message: Good job!


Tests for Q4 (points: 1)
```
>>> lcds([5, 4, 3])
3
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> lcds([4, 5, 4, 4, 3, 2, 1])
4
```
failure_message: Incorrect answer, try again!
success_message: Good job!
```
>>> lcds([4, 5, 6])
1
```
failure_message: Incorrect answer, try again!
success_message: Good job!
